```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Tarea 2: Arreglos y Tablas
Por favor complete este notebook llenando las celdas que se piden. Antes de comenzar, ejecute la celda previa para cargar las pruebas dadas.



```python
# Ejecute esta celda e ignore lo que esta arroja. Esto es solo para verificar que ejecutó la primera celda.
# Si obtuene un NameError, ¡por favor corra la primera celca de este notebook!
grader
```

**Recursos útiles:**
- [Referencia de Python](http://data8.org/fa23/reference/): Hoja de referencia de métodos útiles de matrices y tablas utilizados en fundamentos en ciencia de datos.

**Lecturas recomendadas:**
- [Arreglos](https://inferentialthinking.com/chapters/05/1/Arrays.html)
- [Guía de arreglos de FCD](https://drive.google.com/file/d/1vTduvivHyEX7cdpMCq9B2jC4iy6A3Tru/view?usp=sharing)
- [¿Qué es la ciencia de datos?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causalidad y experimentos](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html)
- [Programando en Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

Para todos los problemas en los que deba escribir explicaciones y oraciones, **deberá** proporcionar la respuesta en el espacio designado. Además, a lo largo de esta tarea y de todas las futuras, ¡asegúrese de no reasignar variables en todo el notebook! Por ejemplo, si usa `temperatura` en su respuesta a una pregunta, no la reasigne más adelante. De lo contrario, ¡reprobará las pruebas que creía que estaba aprobando anteriormente!

**Fecha de entrega:**

Esta tarea deberá ser entregada ** ___ a las ___**. Trabajo que se entregue por fuera del plazo no será recibido.

**Nota: Esta tarea tiene pruebas ocultas. Eso significa que aunque los exámenes indiquen que se aprobó al 100%, no significa que su calificación final será del 100%. Realizaremos más pruebas ocultas de corrección una vez que todos entreguen la tarea.**


Evite compartir respuestas o hacer copia con sus compañeros. Establezca en cambio discusiones grupales con otros estudiantes y docentes del curso.

Si desea contar con el tiempo necesario para pedir ayuda o asesoría en caso de estar atascado, se recomienda iniciar temprano con la realización de las tareas. Tenga en cuenta el cronograma del curso y de entregas.

<!-- BEGIN QUESTION -->

Ingrese su nombre en la siguiente celda:
```python
Nombre = "Marie Curie"
Carnet =  "27182818 "
```

```python
Nombre = "..."
Carnet = "..."
```

<!-- END QUESTION -->

## 1. Creando Arreglos

```python
# Ejecute esta celda para configurar su notebook, pero por favor no la cambie.
import pandas as pd
import numpy as np
import otter
import math
import matplotlib.pylab as plt
from datascience import *
import warnings
warnings.simplefilter('ignore', FutureWarning)
%matplotlib inline
path="."
```

**Pregunta 1.** Cree un arreglo llamado `numeros_extranos` que contenga los siguientes números (en el orden dado) **(4 Points)**:

1. -2
2. floor de 17.6
3. 3
4. 5 a la potencia del ceil de 5.3

*Hint:* `floor` y `ceil` son funciones en el módulo `math`. Importar módulos se trabaja en el Lab 2.

*Nota:* Las listas de Python son y se comportan diferente de los arreglos de NumPy. En este curso usaremos arreglos de NumPy, así que cuando haga un **arreglo**, no haga una lista de Python.


```python
# Nuestra solución puede requier una línea extra de código antes de crear numeros_extraños
...
numeros_extranos = ...
numeros_extranos
numeros_extranos
```

```python
grader.check("q1")
```

**Pregunta 2.** Cree un arreglo llamado `expresiones_sin_comas` que contenga las cinco siguientes cadenas de caracteres: "El juez", "dijo", "el acusado", y  "es inocente" **(1 Puntos)**

```python
expresiones_sin_comas= ...
expresiones_sin_comas
```

```python
grader.check("q2")
```

Las cadenas de caracteres (strings) tienen un método llamado `join`.   

`join` toma como argumento un arreglo de cadenas de caracteres y regresa una cadena de caracteres. Específicamente, el resultado es una sola cadena de caracteres, que es la [concatenación](https://es.wikipedia.org/wiki/Concatenaci%C3%B3n) ("juntar") todas las cadenas de caracteres en el arreglo.


**Pregunta 3.** Use el array `expresiones_sin_comas` y el método `join` para crear dos arreglos **(1 puntos)**:

1. 'El juez, dijo, el acusado, es inocente'(llámelo `usando_comas`, note el espacio después de la coma)
2. 'El juez dijo el acusado es inocente'(llámelo `sin_comas`)



```python
usando_comas = ...
sin_comas = ...
# Estas líneas son solo para imprimir sus respuestas.
print('usando_comas:', usando_comas)
print('sin_comas:', sin_comas)
```

```python
grader.check("q3")
```

## 2. Indexando Arreglos

Estos ejercicios están pensados en dar práctica al momento de acceder a elementos individuales de los arreglos. En Python (y en muchos lenguajes de programación), cada elemento es accesible mediante su  *índice*; por ejemplo, el primer elmento es el elemento con índice 0. Los índices deben ser números **enteros**.


***Nota:* Si usted tiene experiencia previa programando, la notación de corchetes le debe ser familiar. NO use la notación de corchetes cuando esté indexando (i.e. `arr[0]`), pues esto puede conducir a un tipo de dato diferente del que estamos esperando. Le puede además hacer fallar en la revisión.**

En caso de sentirse atascado, revise la [Referencia de Python](http://data8.org/fa23/reference/)


**Pregunta 4.** La siguiente celda crea un arreglo llamado  algunos_numeros. Asigne  a la varible `tercer_elemento`  el tercer elemento de `algunos_numeros`. **(1 Puntos)**


```python
algunos_numeros = make_array(-1, -3, -6, -10, -15)
tercer_elemento = ...
tercer_elemento
tercer_elemento
```

```python
grader.check("q4")
```

**Pregunta 5.** La próxima celda crea una tabla que muestra información sobre los elementos de `algunos_numeros` y su orden. Corra la celda para ver la tabla parcialmente completa y llene la información faltante (las celdas que dicen "Ellipsis") mediante la asignación `espacio_a`, `espacio_b`, `espacio_c`, y `espacio_d` para corregir los elementos en la tabla. **(1 Puntos)**

*Hint* Reemplace los `...` con cadenas de caracteres o números. Recuerde que los índices deben ser **enteros**.


```python
variable_a = ...
variable_b = ...
variable_c = ...
variable_d = ...
elementos_algunos_numeros = Table().with_columns(
    "NombrePosicion", make_array("primera", "segunda", variable_a, variable_b, "quinto"),
    "Indice",                     make_array(variable_c, 1, 2, variable_d, 4),
    "Elemento",                   algunos_numeros)
elementos_algunos_numeros
```

```python
grader.check("q5")
```


**Pregunta 6.** En ocasiones usted querrá encontrar el **último** elemento de un arreglo. Suponga que el arreglo tiene 142 elementos. ¿Cuál es el índice del último elemento? **(1 Puntos)**

*Nota:* Su respuesta deberá ser un número positivo.


```python
indice_ultimo_elemento = ...
indice_ultimo_elemento
```

```python
grader.check("q6")
```

Con frecuencia, no se conoce el número de elementos de un arreglo, su *longitud*. La función `len` toma un solo argumento, por ejemplo, en la siguiente celda leemos la tabla `nacimiento_presidentes` y despues analizamos cuantas filas tiene a traves del metodo `len`

```python
ano_nacimiento_presidentes = Table.read_table(f"{path}/nacimiento_presidentes.csv").column('AnioNacimiento')  #Modificar una vez se ajuste la tabla.
len(ano_nacimiento_presidentes)
```

También se pueden aplicar otas funciones sobre el arreglo, `max`, `min` ...

**Pregunta 7.** La siguiente celda carga un arreglo llamado `anio_nacimiento_presidentes` utilizando el comando `Table.read_table(...).column(...)`, que carga una tabla llamada `nacimiento_presidentes.csv` Este comando devuelve un arreglo con los años de nacimiento de los ex-presidentes de Colombia, específicamente en la columna AnioNacimiento. El año más reciente de nacimiento de ese arreglo,  es el año de nacimiento más grande. Asigne ese año a la variable `anio_nacimiento_mas_reciente` utilizando únicamente métodos de tablas y funciones básicas.

Nota: Evite consultar fuentes externas. Usted debería ser capaz de resolver esta pregunta usando únicamente los métodos provistos por la tabla.

```python
anio_nacimiento_presidentes = Table.read_table("nacimiento_presidentes.csv").column('AnioNacimiento')  
anio_nacimiento_mas_reciente =...
anio_nacimiento_mas_reciente
```

```python
grader.check("q7")
```

**Pregunta 8.** Para finalizar, asigne `min_anio_nacimiento` al mínimo entre el  primer,  decimosexto y ulitmo año de nacimiento  

**Nota:** Use la función `min` de Python y métodos de tabla para encontrar la respuesta. Evite calcular manualmente el resultado, recuerda que el contador en python en empieza en cero. 


```python
min_anio_nacimiento = ...  
min_anio_nacimiento
```

```python
grader.check("q8")
```

## 3. Aritmética Básica de Arreglos

**Pregunta 9.** Multiplique los números 2, 32, 0, y 124 por 64. Asigne cada variable en la siguiente celda de tal forma que  `primer_producto` se asigne al resultado de $2 * 64$, `segundo_producto` al resultado de $32 * 64$ y así sucesivamente. **(1 Puntos)**

*Nota*: Para esta pregunta, **no** use arreglos.

```python
primer_producto = ...
segundo_producto = ...
tercer_producto = ...
cuarto_producto = ...
print("Primer Producto:", primer_producto)
print("Segundo Producto:", segundo_producto)
print("Tercer Producto:", tercer_producto)
print("Cuarto Producto:", cuarto_producto)
```

```python
grader.check("q9")
```

**Pregunta 10.** Ahora, realiza los mismos cálculos, pero esta vez usando un arreglo llamado `numeros`y un solo operador de multiplicación (`*`). Almacene los 4 resultados en un arreglo llamado `productos`. **(1 Puntos)**

```python
numeros = ...
productos = ...
productos
productos
```

```python
grader.check("q10")
```

**Pregunta 11.** Ups, cometimos un error. En lugar de 64, queremos que multiplique cada número por 128. Calcule los productos correctos en la celda de abajo usando aritmética de arreglos. Note que su trabajo es realmente sencillo si previamente definió un arreglo que contiene los 4 números. **(1 Puntos)**

```python
productos_correctos = ...
productos_correctos
```

```python
grader.check("q11")
```

**Pregunta 12.** Hemos cargado un arreglo de la temperatura promedio por año de la ciudad de Bogotá y Barcelona, desde el año 1995,   todas las temperaturas están en Fahrenheit. Convierta todas las temperaturas a Celsius:

$C = (F - 32) * \frac{5}{9}$. 

llame esta variable `celsius_temps`. Ahora, asegúrese de **REDONDEAR** el resultado final al entero más cercano usando la función `np.round` llame esta variable `celsius_temps_redondeadas` **(1 Puntos)**

```python
temperatura = Table().read_table("TemperaturaBogotaBarcelona.csv").column("AvgTemperature")
celsius_temps = ...
celsius_temps_redondeadas = ...
```

```python
grader.check("q12")
```

## 4. Old Faithful

*Old* Faithful es un [géiser](https://es.wikipedia.org/wiki/G%C3%A9iser) en Yellowstone que erupciona cada 44-125 minutos (de acuerdo con [Wikipedia](https://es.wikipedia.org/wiki/Old_Faithful#:~:text=El%20Old%20Faithful%2C%20traducido%20al,%2C%20en%20Wyoming%2C%20Estados%20Unidos)). A las personas se les dice con frecuencia que erupciona cada una hora, pero de hecho el tiempo de espera entre erupciones es más variable. Veámoslo.

**Pregunta 13.** La primera línea de la celda de abajo asigna `tiempos_espera` a un arreglo con 272 tiempos de espera consecutivos entre erupciones, tomados de un conjunto de datos de 1938. Asigne los nombres `menor`, `mayor` y `promedio`de forma tal que la asignación `print` esté correcta. **(1 Puntos)**, emplee los metodos `np.min`, `np.max`, `np.mean`

```python
tiempos_espera = Table.read_table('old_faithful.csv').column('espera')

menor = ...
mayor = ...
promedio = ...
print("Old Faithful erupciona entre", menor, "a", mayor, "minutos y cada", promedio, "minutos en promedio.")
```

```python
grader.check("q13")
```

**Pregunta 14.** Asigne a la variable `mayor_diferencia` el  mayor cambio en la espera entre dos erupciones consecutivas. Por ejemplo, la tercera erupción ocurrió luego de 74 minutos y la cuarta luego de 62 minutos, así que la diferencia fue 74 - 62 = 12 minutos. **(1 Puntos)**

*Hint*: Queremos el valor absoluto del valor de la mayor diferencia.

*Nota*: Emplea `np.diff()` y calcula la diferencia entre valores concecutivos, aplica el valor absoluto y después determina el `max` 

```python
mayor_diferencia=....
mayor_diferencia
```

```python
grader.check("q14")
```

**Pregunta 15.** La tabla `numero_erupciones_espera` contiene dos columnas: `numero_erupcion`, que representa el número de cada erupción, y `espera`, que indica el tiempo (en minutos) que se tuvo que esperar desde el final de esa erupción hasta que ocurrió la siguiente. Por ejemplo, considere las primeras dos filas de la tabla:

| NumeroErupcion | Espera  |
|--------------------|---------|
| 1                  | 79      |
| 2                  | 54      |

Esto se puede interpretar de la siguiente manera: después de la primera erupción, se espera 79 minutos hasta que ocurre la segunda erupción. Luego, después de la segunda erupción, se espera 54 minutos para la tercera erupción.

Ahora, suponga que Oscar y Wendy comenzaron a observar Old Faithful al inicio de la primera erupción. Ellos continúan observando hasta que finaliza la décima erupción.

La pregunta es: ¿Cuántos minutos en total pasarán esperando entre erupciones? **(1 Punto)**

*Hint #1:* Puede usar el método `take` en la tabla `numero_erupciones_espera` para seleccionar las filas necesarias.

*Hint #2:* `tiempo_espera_primeras_nueve` deberá ser un arreglo que contenga los tiempos de espera entre las primeras nueve erupciones.


```python
tabla = Table.read_table('old_faithful.csv')
tabla = tabla.with_column('NumeroErupcion', range(1, tabla.num_rows + 1))
numero_erupciones_espera = tabla.select("NumeroErupcion", "espera")
espera = ...
tiempo_espera_primeras_nueve = ...
tiempo_espera_primeras_nueve
```

```python
grader.check("q15")
```

**Pregunta 16.** Imaginemos que su estimación del siguiente tiempo de espera siempre fue la duración del tiempo de espera anterior. Si siempre adivinara el tiempo de espera anterior, ¿qué tan grande sería, en promedio, el error al adivinar los tiempos de espera?. Emplee toda la tabla: 

Por ejemplo, dado que los primeros cuatro tiempos de espera son 79, 54, 74 y 62, la diferencia promedio entre su "predicción" y el tiempo real solo para la segunda, tercera y cuarta erupciones sería: $\frac{|79-54|+ |54-74|+ |74-62|}{3} = 19$.

En la variable `error` cálcule el promedio.

Ayuda: Puede ser de ayuda emplear los metodos  `abs`, `np.sum`, `len` o `np.mean`

```python
tabla = Table.read_table('old_faithful.csv')
tabla = tabla.with_column('NumeroErupcion', range(1, tabla.num_rows + 1))
numero_erupciones_espera = tabla.select("NumeroErupcion", "espera")
espera = numero_erupciones_espera.column("espera")

diferencias = np.diff(tiempos_espera)
error = ...
error
```

```python
grader.check("q16")
```

## 5. Tablas

**Pregunta 17.** Suponga que tiene 4 manzanas y 3 piñas. (Quizás esté usando Python para resolver un problema de álgebra del bachillerato). Cree una tabla que contenga esta información, esta deberá contener dos columnas: `nombre_de_la_fruta` y `cantidad`. Asigne la nueva tabla a la variable `frutas`. **(1 Puntos)**

*Note:* Use letras minúsculas y una sola palabra para el nombre de cada fruta, por ejemplo `"manzana"`.




|nombre_de_la_fruta|cantidad|
|------------------|-----|
|manzana|4|
|pina|3|


```python
frutas = Table().with_columns( ...)
frutas
```

```python
grader.check("q17")
```

**Pregunta 18.** El archivo `inventario.csv` contiene información sobre el inventario de un puesto de frutas. Cada fila representa el contenido de una caja de fruta. Cargue el archivo como una tabla llamada `inventario` usando la función `Table.read_table()`. `Table.read_table(...)` toma un argumento (el nombre de un archivo en formato cadena de caracteres) y regresa una tabla. **(1 Puntos)**

Nota: Un archivo csv puede estar seprado por comas "," , punto y coma ";"  y pip "|", en este caso,  esta separado por ";", por ello debes indicar como esta separadas las columnas del archivo, 

```python
inventario = Table().read_table(nombre_archivo, sep=";")

```

```python
inventario = 
inventario
```

```python
grader.check("q18")
```

**Pregunta 19.** Considere los ID de las cajas y las frutas que contienen. ¿Existen cajas con ID diferentes que contengan frutas del mismo tipo?. Asigne a una nueva columna llamada  `todas_diferentes` dentro de la tabla `inventario`con el valor `True` si **todas** las cajas contienen frutas diferentes. Si alguna fruta aparece en más de una caja (aunque los ID de las cajas sean diferentes), asigne `False` a `todas_diferentes`. **(1 Punto)**

*Hint:* No es necesario escribir un código para calcular el valor de `todas_diferentes`. Simplemente observe la tabla y asigne el valor de `True` o `False` de acuerdo a las frutas que se repiten.

```python
inventario.with_column('nueva_columna', [valor1, valor2, ...])
```

```python
inventario=...
inventario
```

```python
grader.check("q19")
```

**Pregunta 20.** El archivo `ventas.csv` contiene el número de frutas vendidas de cada caja el sábado pasado. Esta tiene una columna extra llamada `PrecioPorFruta`, la cual es el precio *por unidad*  por fruta en dicha caja. Las filas están en el mismo orden que en la tabla `inventario`.Cargue estos datos en una tabla llamada `ventas`. **(5 Puntos)**

```python
ventas=...
ventas
```

```python
grader.check("q20")
```

**Pregunta 22.** ¿Cuántas frutas se vendieron en total en la tienda ese día? **(1 Puntos)**, guarde el resultado en una variable llamada `total_frutas_vendidas`


```python
total_frutas_vendidas=...
total_frutas_vendidas
```

```python
grader.check("q22")
```

**Pregunta 21.** ¿Cuál fue la ganancia total de la tienda (el precio total de todas las frutas vendidas) **(1 Puntos)** guarde el resultado en una variable llamada `ganancia_total`


*Hint:* Si se atasca, piense primero en cómo podría calcular la ganancia total solo de las ventas de uvas.




```python
ganancia_total=...
ganancia_total
```

```python
grader.check("q21")
```


---

¡Terminaste con la tarea 2!

**Información importante sobre el envío:**
- **Ejecuta todas las pruebas** y verifica que todas pasan
- **Guardar** desde el menú **Archivo**
- **Ejecute la celda final **
- Luego, ve a este [enlace](https://cienciadatosudea.github.io/001_SitioWebFundCienciaDatos/04entregas.html) y envía el archivo ipynb a la tarea correspondiente. El nombre de esta tarea es "Lab XX Autograder", donde XX es el número de laboratorio: 01, 02, 03, etc.

- Si terminas temprano en el laboratorio regular, **registra tu asistencia con el instructor**.

**Es su responsabilidad asegurarse de que su trabajo esté guardado antes de ejecutar la última celda.**

`Referencia`: Este material esta basado en [Data 8: The Foundations of Data Science](https://www.data8.org/) - University of California, Berkeley.


## Estudiantes de Fundamentos de Ciencia de Datos

¡Felicitaciones por terminar la tarea 2!

<img src="https://raw.githubusercontent.com/CienciaDatosUdea/001_DataSetFundCienciaDatos/refs/heads/main/semestre2024-2/hw/hw02/kim.jpg" width="30%" alt="Close up picture of dog smiling."/>

Mascota de la semana: **Kim**

## Entrega

Asegúrese que ha ejecutado todas las celdas en su notebook en orden antes de ejecutar la celda de abajo. De esta forma todas las imágenes/gráficas apareceran en la salida. La celda de abajo va a generar un archivo con extensión .zip para que usted lo entregue. **¡Por favor guarde antes de exportar!**


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert numeros_extranos.size == 4
>>> esperado = make_array(-2 * np.pi, 17 * np.pi, 3 * np.pi, 15625 * np.pi)
>>> for i in range(numeros_extranos.size):
...     valor_real = numeros_extranos.item(i)
...     valor_esperado = esperado.item(i)
...     assert np.isclose(valor_real * np.pi, valor_esperado, 0.01)
```


Tests for q10 (points: 1)
```
>>> def test_productos1():
...     numeros = make_array(2, 32, 0, 124)
...     factor = 64
...     productos_esperados = factor * numeros
...     assert all(productos == productos_esperados)
>>> test_productos1()
```


Tests for q11 (points: 1)
```
>>> def test_productos2():
...     numeros = make_array(2, 32, 0, 124)
...     factor = 128
...     productos_esperados = factor * numeros
...     assert all(productos_correctos == productos_esperados)
>>> test_productos2()
```


Tests for q12 (points: 1)
```
>>> value_a = 14.596153846153847
>>> value_b = 14.566698067200857
>>> assert np.isclose(np.mean(celsius_temps_redondeadas), value_a, 0.001)
>>> assert np.isclose(np.mean(celsius_temps), value_b, 0.001)
```


Tests for q13 (points: 1)
```
>>> assert np.isclose(menor, 7, 0.001)
>>> assert np.isclose(mayor, 96, 0.001)
>>> assert np.isclose(promedio, 70.6654411765, 0.001)
```


Tests for q14 (points: 1)
```
>>> assert np.isclose(mayor_diferencia, 74, 0.01)
```


Tests for q15 (points: 1)
```
>>> assert np.isclose(tiempo_espera_primeras_nueve / np.pi, 633 / np.pi, 0.001)
```


Tests for q16 (points: 1)
```
>>> assert np.isclose(error / np.pi, 20.908 / np.pi, 0.01)
```


Tests for q17 (points: 1)
```
>>> def test_frutas_table():
...     assert 'nombre_de_la_fruta' in frutas.labels
...     promedio_cantidad = np.mean(frutas.column('cantidad'))
...     assert promedio_cantidad == 3.5
>>> test_frutas_table()
```


Tests for q18 (points: 1)
```
>>> assert inventario.labels == ('ID de caja', 'nombre de la fruta', 'cantidad'), 'Las columnas no coinciden con las esperadas.'
>>> assert inventario.num_rows == 8
>>> assert inventario.column('nombre de la fruta').item(0) == 'kiwi', "El primer valor de la columna 'nombre de la fruta' debería ser 'kiwi'."
>>> assert inventario.column('cantidad').item(0) == 45, "El primer valor de la columna 'cantidad' debería ser 45."
>>> assert inventario.column('nombre de la fruta').item(1) == 'fresa', "El segundo valor de la columna 'nombre de la fruta' debería ser 'fresa'."
>>> assert inventario.column('cantidad').item(1) == 123, "El segundo valor de la columna 'cantidad' debería ser 123."
```


Tests for q19 (points: 1)
```
>>> valores_esperados = [True, False, True, True, False, True, False, True]
>>> assert inventario.column('todas_diferentes').tolist() == valores_esperados
```


Tests for q2 (points: 1)
```
>>> test_para_ocultar1 = 'es acusado testimonio testigo culpable El inocente fiscal tribunal jurado juez el dijo'
>>> test_para_ocultar2 = 'es acusado testimonio testigo culpable El inocente fiscal tribunal jurado juez el dijo'
>>> esperado = test_para_ocultar1 + ''.join(make_array('El juez', 'dijo', 'el acusado', 'es inocente')) + test_para_ocultar1
>>> expresiones_ = test_para_ocultar1 + ''.join(expresiones_sin_comas) + test_para_ocultar1
>>> assert expresiones_sin_comas.size == 4, 'Error: el arreglo no tiene 7 elementos.'
>>> assert expresiones_ == esperado
```


Tests for q20 (points: 1)
```
>>> columnas_actuales = [col.strip() for col in ventas.labels]
>>> columnas_esperadas = ['ID de caja', 'nombre de la fruta', 'cantidad', 'CantidadVendida', 'PrecioPorFruta']
>>> for columna in columnas_esperadas:
...     assert columna in columnas_actuales, f"La columna '{columna}' no existe en la tabla."
```


Tests for q21 (points: 1)
```
>>> v = 229700
>>> assert ganancia_total == v
```


Tests for q22 (points: 1)
```
>>> value1 = 638
>>> assert total_frutas_vendidas / np.pi == value1 / np.pi
```


Tests for q3 (points: 1)
```
>>> assert len(sin_comas) == 35
>>> assert len(usando_comas) == 38
```


Tests for q4 (points: 1)
```
>>> ae = -6
>>> assert tercer_elemento == ae
```


Tests for q5 (points: 1)
```
>>> def test_elementos_algunos_numeros():
...     assert elementos_algunos_numeros.num_columns == 3
...     suma_nombres_columnas = sum([sum((ord(c) for c in col)) for col in elementos_algunos_numeros.labels])
...     assert suma_nombres_columnas == 2860
...     assert elementos_algunos_numeros.num_rows == 5
...     suma_indices = sum(elementos_algunos_numeros.column('Indice'))
...     assert suma_indices == 10
...     suma_elementos = sum(elementos_algunos_numeros.column('Elemento'))
...     assert suma_elementos == -35
>>> test_elementos_algunos_numeros()
```


Tests for q6 (points: 1)
```
>>> value = 141
>>> assert indice_ultimo_elemento / np.pi == value / np.pi
```


Tests for q7 (points: 1)
```
>>> assert isinstance(anio_nacimiento_mas_reciente, (int, np.integer))
>>> assert anio_nacimiento_mas_reciente + 10 == 1976 + 10
```


Tests for q8 (points: 1)
```
>>> assert np.isclose(min_anio_nacimiento, 1783, 0.001)
```


Tests for q9 (points: 1)
```
>>> def test_productos():
...     b = 64
...     primer_producto_esperado = 128
...     segundo_producto_esperado = 2048
...     tercer_producto_esperado = 0
...     cuarto_producto_esperado = 7936
...     assert primer_producto == primer_producto_esperado
...     assert segundo_producto == segundo_producto_esperado
...     assert tercer_producto == tercer_producto_esperado
...     assert cuarto_producto == cuarto_producto_esperado
>>> test_productos()
```
