```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

# Tarea 3: Manipulación de Tablas y Visualización


Por favor complete este notebook llenando las celdas que se indican. Antes de comenzar, ejecute la celda previa para cargar las pruebas.


**Recursos útiles:**
- [Referencia de Python](http://data8.org/fa23/reference/): Hoja de referencia y métodos útiles de métodos de arreglos y tablas.

**Lectura recomendada**:
* [Visualización](https://inferentialthinking.com/chapters/07/Visualization.html)
* [Histogramas](https://inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html)

Para todos los problemas en los que deba escribir explicaciones y oraciones, **deberá** proporcionar la respuesta en el espacio designado. Además, a lo largo de esta tarea y de todas las futuras, ¡asegúrese de no reasignar variables en todo el notebook! Por ejemplo, si usa `temperatura` en su respuesta a una pregunta, no la reasigne más adelante. De lo contrario, ¡reprobará las pruebas que creía que estaba aprobando anteriormente!

**Fecha de entrega:**

Esta tarea deberá ser entregada ** ___ a las ___**. Si quiere puntos extra, entréguela el (un día antes). Trabajo que se entregue por fuera del plazo no será recibido.

**Nota: Esta tarea tiene pruebas ocultas. Eso significa que aunque los exámenes indiquen que se aprobó al 100%, no significa que su calificación final será del 100%. Realizaremos más pruebas ocultas de corrección una vez que todos entreguen la tarea.**


Evite compartir respuestas o hacer copia con sus compañeros. Establezca en cambio discusiones grupales con otros estudiantes y docentes del curso.

Si desea contar con el tiempo necesario para pedir ayuda o asesoría en caso de estar atascado, se recomienda iniciar temprano con la realización de las tareas. Tenga en cuenta el cronograma del curso y de entregas.

```python
# No cambie esta celda; solo ejecútela.
import numpy as np
from datascience import *
import warnings
import matplotlib.pyplot as plt
%matplotlib inline
plt.style.use('fivethirtyeight')
warnings.simplefilter('ignore', FutureWarning)
```

<!-- BEGIN QUESTION -->

Ingrese su nombre en la siguiente celda:
```python
Nombre = "Ada Lovelace"
Carnet =  "3,141592...."
```

```python
Nombre = "..."
Carnet = "..."
```

<!-- END QUESTION -->

## 1. Datos de Uber

**Nota:** Recomendamos leer el [capítulo 7.2](https://inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html) del libro guía antes de llegar a la pregunta 3.

Las siguientes tablas contienen datos de 200000 viajes de Uber realizados durante una semana en las áreas metropolitanas de Manila y Boston. Estos datos provienen del proyecto [Uber Movement](https://movement.uber.com), con la siguiente información:

- La columna `sourceid` contiene los códigos de los lugares de inicio de los viajes.
- La columna `dstid` incluye los códigos de los destinos correspondientes.
- La columna `hod` indica la hora del día en la que se realizó cada viaje (en formato de código).
- La columna `ride time` representa la duración del viaje en minutos.

<!-- La siguiente tabla contiene datos de diferentes viajes de uber en la ciudad de bogotá, datos tomados de [kaggle](https://www.kaggle.com/datasets/ishandutta/uber-travel-movement-data-2-billion-trips).  Las columnas `fuente` y `destino` contienen coordenadas geográficas de los lugares de inicio y de destino de cada viaje. La columna `` contiene códigos que corresponden a la hora del día en la que el viaje tuvo lugar. La columna `ride time` contiene la duración del viaje en minutos. -->

```python
boston = Table.read_table("boston.csv")
manila = Table.read_table("manila.csv")
print("Tabla de Boston")
boston.show(4)
print("Tabla de Manila")
manila.show(4)
```

<!-- BEGIN QUESTION -->

**Pregunta 1.** Genere un histograma que permita visualizar la distribución de los tiempos de viaje en Boston usando los bins dados en `bins_iguales`. **(1 Puntos)**

*Hint:* Vea el [capítulo 7.2](https://inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html) si no sabe cómo especificar los bins.

```python

# Leea el archivo boston.csv
bins_iguales = np.arange(0, 120, 5)
...

# ¡No borre la siguiente línea!
plt.ylim(0, 0.05);
```

```python
grader.check("q1")
```



**Pregunta 2.** Genere un histograma que muestre la distribución de los tiempos de viaje registrados en Manila, utilizando los intervalos de clase (bines) proporcionados. (1 punto)

```python

# Leea el archivo manila.csv
bins_iguales = np.arange(0, 120, 5)
...

# ¡No borre la siguiente línea!
plt.ylim(0, 0.05);
```

```python
grader.check("q2")
```


**Pregunta 3.** Observa con más detalle la etiqueta del eje y en los gráficos anteriores. Asigna a la variable `unit_meaning` un número entero (1, 2, 3) que represente correctamente la "unidad" mencionada en la frase "Porcentaje por unidad". **(1 Punto)**

1. minuto  
2. tiempo de viaje, (`ride time`)
)  
3. segundos


```python
unit_meaning = ...
```

```python
grader.check("q3")
```

**Pregunta 4.** Asigna a las variables `boston_menor_15` y `manila_menor_15` el **porcentaje de viajes con duración menor a 15 minutos** en las respectivas áreas metropolitanas de Boston y Manila.

Para resolverlo, **utiliza únicamente las variables de altura que se proporcionan más adelante**, junto con operaciones matemáticas básicas.  
No debes inspeccionar visualmente las tablas `boston` ni `manila`; en cambio, implementa la solución computacionalmente, filtrando los datos según corresponda.

**(1 punto)**



Ayuda: 
1. Revisa el ancho de cada bin 
2. Con las alturas dada más abajo, calcule el área

El área de cada bin es: 
\begin{equation}
A = \mathrm{Altura\ bin} \times  \mathrm{Ancho\  bin}
\end{equation}
3. Cálcula el porcentaje de viajes para cada altura
4. ¿Cuál es el porcentaje de viajes que tuvieron una duración menor a 15 minutos?

```python
boston_menor_5_bin_altura = 1.2
manila_menor_5_bin_altura = 0.6
boston_5_a_menor_10_bin_altura = 3.2
manila_5_a_menor_10_bin_altura = 1.4
boston_10_a_menor_15_bin_altura = 4.9
manila_10_a_menor_15_bin_altura = 2.2
boston_menor_15 = ...
manila_menor_15 = ...
```

```python

```

```python
grader.check("q4")
```

**Pregunta 5.** Analicemos más de cerca la distribución de los tiempos de viaje en Boston. Asigna la variable `boston_mediana_bin` a un número entero (1, 2, 3 o 4) que corresponda al bin que contiene la mediana del tiempo de viaje. **(1 Puntos)**

1. 0-8 minutos  
2. 8-14 minutos  
3. 14-20 minutos  
4. 20-40 minutos  

Ayuda: La mediana es una lista ordenada de números,  que divide a la lista en dos partes iguales, de manera que la mitad de los elementos se encuentran por debajo de la mediana y la otra mitad por encima.



```python
boston_mediana_bin = ...
boston_mediana_bin
```

```python
grader.check("q5")
```

<!-- BEGIN QUESTION -->

La desviación estandar  describe cómo de dispersos o concentrados están los datos alrededor de la media o promedio, si la desviación estándar es pequeña, los datos están más concentrados cerca de la media (más consistentes). Si la desviación estándar es grande, los datos están más dispersos y se extienden más lejos de la media (más variabilidad).

Ejemplo: 

Fórmula para calcular la media (o promedio):

\begin{equation}
\text{Media} = \frac{1}{n} \sum_i x_i
\end{equation}

Donde:
- $x_i$ son los valores individuales de los datos.
- $n$ es el número total de datos.

Por ejemplo, si tienes los números $5, 10, 15$, la media o promedio sería:

\begin{equation}
\text{Media} = \frac{5 + 10 + 15}{3} = 10
\end{equation}


Fórmula para calcular la desviación estándar:

La desviación estándar (\(\sigma\)) mide cuánto se desvían los datos respecto a la media. La fórmula es la siguiente:

\begin{equation}
\sigma = \sqrt{\frac{1}{n} \sum_{i=1}^{n} (x_i - \mu)^2}
\end{equation}

Donde:
- $\sigma$ es la desviación estándar.
- $n$ es el número total de datos.
- $x_i$ son los valores individuales de los datos.
- $\mu$ es la media de los datos.

### Ejemplo:

Supongamos que tenemos los valores: \(5, 10, 15\).

1. **Calcular la media**:
\begin{equation}
   \mu = \frac{5 + 10 + 15}{3} = 10
\end{equation}

2. **Restar la media a cada valor (diferencias)**:
   - $5 - 10 = -5$
   - $10 - 10 = 0$
   - $15 - 10 = 5$

3. **Elevar al cuadrado cada diferencia**:
   - $(-5)^2 = 25$
   - $0^2 = 0$
   - $5^2 = 25$

4. **Sumar los cuadrados de las diferencias**:
\begin{equation}
   25 + 0 + 25 = 50
\end{equation}

5. **Dividir por el número total de datos** (\(n = 3\)):
\begin{equation}
   \frac{50}{3} \approx 16.67
\end{equation}

6. **Sacar la raíz cuadrada** para obtener la desviación estándar:
\begin{equation}
   \sigma = \sqrt{16.67} \approx 4.08
\end{equation}





**Pregunta 6.** Identifique una diferencia entre los histogramas, en términos de sus propiedades estadísticas. ¿Qué puede decir sobre el promedio y/o la desviación de cada histograma? **(1 Puntos)**

*Hint*: La mejor forma de hacerlo es comparar los dos histogramas (de los puntos 1 and 2), no tiene que hacer ningún cálculo, solo realice la comparación  visualmente.


```python
_Escriba aquí su respuesta, reemplazando este texto._
```

<!-- BEGIN QUESTION -->

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Pregunta 7.** ¿Por qué su respuesta en la pregunta 6 es correcta? Basado en las siguientes lecturas, ¿por qué las distribuciones de Boston y Manila son diferentes? **(4 Puntos)**

- [Lectura caso Boston](https://www.climatestotravel.com/climate/united-states/boston)
- [Lectura caso Manila](https://manilafyi.com/why-is-manila-traffic-so-bad/)

*Hint:* Trate de pensar en factores externos de ambas ciudades que pueden estar causando dichas diferencias. Las lecturas le dan algunos facvtores potenciales -- Trate de conectarlos con los datos de tiempos de viaje.



```python
Escriba aquí su respuesta, reemplazando este texto._

```

<!-- END QUESTION -->

<!-- END QUESTION -->

## 4. Histograms

Considere la siguiente gráfica de dispersión:

![Alt text](scatter.png "Scatter plot showing data points for the variables 'x' and 'y'. The data are symmetric about the x-axis centered at 0 and symmetric about the y-axis centered at 0, but with no data in the [-0.5, 0.5] range on the y-axis.")

Los ejes de la gráfica representan valores de dos variables: $x$ y $y$.


Suponga que tiene una tabla llamada `t` que contiene dos columnas en ella

- `x`: una columna que contiene los valores de x de los puntos en la gráfica de dispersión.
- `y`: una columna que contiene los valores de y de los puntos en la gráfica de dispersión.

Abajo se  dan tres histogramas. Uno corresponde a la columna `x`, uno corresponde a la columna `y`, y otro a ninguna de las dos columnas.

**Histogram A:**

![Alt text](var3.png "Symmetrical, bell-shaped histogram centered around 0")

**Histogram B:**

![Alt text](var1.png "Symmetrical histogram with two peaks at -1 and 1 but no data around 0")

**Histogram C:**

![Alt text](var2.png "Asymmetrical histogram with a peak around -0.5 and a right skew")

**Pregunta 8.** Suponga que usted corre `t.hist('x')`. ¿Cuál histograma debería producir este código `t.hist('x')`? Asigne `histograma_columna_x` a 1, 2, ó 3. **(5 Puntos)**

1. Histograma A
2. Histograma B
3. Histograma C


```python
histograma_columna_x = ...
```

```python
grader.check("q8")
```

<!-- BEGIN QUESTION -->


**Pregunta 9.** Dé al menos una razón por la cual eligió el histograma de la Pregunta 8. **Asegúrese de indicar claramente cuál histograma eligió** (ejemplo: "Yo elegí el histograma A porque..."). 


```python
Escriba su respuesta aqui...reemplazando este texto
```

<!-- END QUESTION -->


**Pregunta 10.** Suponga que usted corre `t.hist('y')`. ¿Cuál histograma debería producir este código? Asigne `histograma_columna_y` a 1, 2, ó 3. **(5 Puntos)**

1. Histograma A
2. Histograma B
3. Histograma C



```python
histograma_columna_y = 2

```

```python
grader.check("q10")
```

<!-- BEGIN QUESTION -->

<!-- BEGIN QUESTION -->

**Pregunta 11.** Dé al menos una razón por la cual eligió el histograma de la Pregunta 3. **Asegúrese de indicar claramente cuál histograma eligió** (ejemplo: "Yo elegí el histograma A porque..."). 


```python
Escriba su respuesta aqui...reemplazando este texto
```

<!-- END QUESTION -->

_Escriba aquí su respuesta, reemplazando este texto._

¡Terminó la tarea 3!

**Información sobre la entrega:** Asegúrese de correr todas las pruebas y verificar que se ejectuen correctamente. Aquí, seleccione **Guardar Notebook** En el menú de **Archivo** y **corra la  última celda**. Descargue el archivo .zip y el archivo con extensión .ipynb y subálo al enlace compartido en la página del curso para su entrega. El nombre de esta tarea es "HW 03 ". **Es su responsabilidad asegurarse que su trabajo se guarde antes de correr la última celda**.



`Referencia`: Este material esta basado en [Data 8: The Foundations of Data Science](https://www.data8.org/) - University of California, Berkeley.

## Estudiantes de Fundamentos de Ciencia de Datos

¡Felicitaciones por terminar la tarea 3!

<img src="./lego.png" width="30%" alt="Close up picture of dog smiling."/>

Perro de la semana: **Lego**



## Entrega

Asegúrese que ha ejecutado todas las celdas en su notebook en orden antes de ejecutar la celda de abajo. De esta forma todas las imágenes/gráficas apareceran en la salida. La celda de abajo va a generar un archivo con extensión .zip para que usted lo entregue. **¡Por favor guarde antes de exportar!**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert 'ride time' in boston.labels, "La columna 'ride_time' no existe en la tabla."
>>> assert len(boston.column('ride time')) > 0, "La columna 'ride_time' está vacía."
>>> assert np.isclose(np.mean(boston.column('ride time')), 16.502477909999996, 0.001)
>>> assert len(bins_iguales) - 1 == len(np.histogram(boston.column('ride time'), bins=bins_iguales)[0]), 'El número de bins no coincide con los bins proporcionados.'
```


Tests for q10 (points: 1)
```
>>> assert histograma_columna_y == 2
```


Tests for q2 (points: 1)
```
>>> assert 'ride time' in manila.labels, "La columna 'ride_time' no existe en la tabla."
>>> assert len(manila.column('ride time')) > 0, "La columna 'ride_time' está vacía."
>>> assert np.isclose(np.mean(manila.column('ride time')), 33.74324175, 0.001)
>>> assert len(bins_iguales) - 1 == len(np.histogram(manila.column('ride time'), bins=bins_iguales)[0]), 'El número de bins no coincide con los bins proporcionados.'
```


Tests for q3 (points: 1)
```
>>> assert unit_meaning == 2
```


Tests for q4 (points: 1)
```
>>> assert np.isclose(manila_menor_15, 21, 0.001)
>>> assert np.isclose(boston_menor_15, 46.5, 0.001)
```


Tests for q5 (points: 1)
```
>>> assert boston_mediana_bin == 3
```


Tests for q8 (points: 1)
```
>>> assert histograma_columna_x == 3
```
