```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Tarea 5: Funcion apply, iteraciones y condicionales

Tarea 4 no tuvimos, por eso saltamos a la 5.

Por favor complete este notebook llenando las celdas que se indican. 

Antes de comenzar, ejecute la celda previa para cargar las pruebas.


**Recursos útiles:**
- [Referencia de Python](http://data8.org/fa23/reference/): Hoja de referencia y métodos útiles de métodos de arreglos y tablas.


Para todos los problemas en los que deba escribir explicaciones y oraciones, **deberá** proporcionar la respuesta en el espacio designado. Además, a lo largo de esta tarea y de todas las futuras, ¡asegúrese de no reasignar variables en todo el notebook! Por ejemplo, si usa `temperatura` en su respuesta a una pregunta, no la reasigne más adelante. De lo contrario, ¡reprobará las pruebas que creía que estaba aprobando anteriormente!

**Fecha de entrega:**

Esta tarea deberá ser entregada ** ___ a las ___**. Si quiere puntos extra, entréguela el (un día antes). Trabajo que se entregue por fuera del plazo no será recibido.

**Nota: Esta tarea tiene pruebas ocultas. Eso significa que aunque los exámenes indiquen que se aprobó al 100%, no significa que su calificación final será del 100%. Realizaremos más pruebas ocultas de corrección una vez que todos entreguen la tarea.**


Evite compartir respuestas o hacer copia con sus compañeros. Establezca en cambio discusiones grupales con otros estudiantes y docentes del curso.

Si desea contar con el tiempo necesario para pedir ayuda o asesoría en caso de estar atascado, se recomienda iniciar temprano con la realización de las tareas. Tenga en cuenta el cronograma del curso y de entregas.

```python
# No cambie esta celda; solo ejecútela.
import numpy as np
from datascience import *
import warnings
import matplotlib.pyplot as plt
%matplotlib inline
plt.style.use('fivethirtyeight')
warnings.simplefilter('ignore', FutureWarning)
```

<!-- BEGIN QUESTION -->

Ingrese su nombre en la siguiente celda:
```python
Nombre = "Ada Lovelace"
Carnet =  "3,141592...."
```

```python
Nombre = "..."
Carnet = "..."
```

<!-- END QUESTION -->

## 1. Datos de cafe colombia

**Pregunta 1** Leer la tabla cafe_colombia, la cual tiene datos que  corresponden a las **Evaluaciones Agropecuarias Municipales de Colombia**, provenientes de fuentes de datos abiertos de colombia. Contiene estadísticas anuales sobre cultivos y producción agrícola en diferentes municipios del país.

Las  Columnas del conjunto de datos son:

- **departamento**: Nombre del departamento.
- **municipio**: Nombre del municipio.
- **año**: Año de recolección del dato.
- **periodo**: Período o ciclo del cultivo (usualmente anual o semestral).
- **areaSembradaHa**: Superficie sembrada en hectáreas.
- **areaCosechadaHa**: Superficie efectivamente cosechada en hectáreas.
- **produccionT**: Producción obtenida en toneladas.
- **cicloDeCultivo**: Ciclo de duración del cultivo (ej. permanente, transitorio).

> 🔍 Antes de procesar, verifica que los valores numéricos estén correctamente tipificados y sin caracteres especiales (como comas o saltos de línea).


```python
cafe = ....
```

```python
grader.check("q1")
```

**Pregunta 2** Define una función ```calcular_rendimiento``` que reciba como argumentos la columnas asociadas a la producción y el área cosechada, y retorne el rendimiento como: 

\begin{equation}
\text{rendimiento} = \frac{\text{producción}} {\text{área.}}
\end{equation}



Luego crea una nueva tabla llamada cafe_rendimiento

```python
def calcular_rendimiento(produccion, area):
    if area==0:   
        return np.nan     # Que pasa si no incluyes este condicional
    
    return ....

cafe_rendimiento_values = cafe.apply(calcular_rendimiento, ['...', '...'])
cafe_rendimiento=...
cafe_rendimiento
```

```python
grader.check("q2")
```

**Pregunta 3**
Escribe una función llamada `produccion_departamento` que reciba:

- `cafe_rendimiento`: una tabla con datos históricos de producción de café. Esta tabla incluye, al menos, las columnas `'departamento'`, `'año'` y `'produccionT'`.
- `departamento`: el nombre de un departamento como cadena de texto.

La función debe:

1. Filtrar la tabla para conservar solo las filas correspondientes al departamento especificado.
2. Seleccionar únicamente las columnas `'año'` y `'produccionT'`.
3. Agrupar los datos por año, sumando la producción total por cada año.
4. Renombrar las columnas a `'anio'` y `'produccionTotal'`.
5. Ordenar los resultados de mayor a menor según la producción total.

- Salida esperada
Una tabla con dos columnas:
- `anio`: el año de producción.
- `produccionTotal`: la suma de la producción total (en toneladas) para ese año.

-  Ejemplo de uso (sin código)
```python
produccion_departamento(cafe, "Antioquia")

```python
def produccion_departamento(cafe_rendimiento, departamento):
    # Filtrar por el departamento dado
    cafe_dep = cafe_rendimiento.where('departamento', departamento)
    
    # Agrupar por año y sumar la producción
    produccion = (
        cafe_dep.select('año', 'produccionT')
        .group('año', sum)
        .relabeled('año', 'anio')
        .relabeled('produccionT sum', 'produccionTotal')
        .sort('produccionTotal', descending=True)
    )
    return produccion
produccion_departamento(cafe_rendimiento,"ANTIOQUIA")
```

```python
grader.check("q3")
```

**Pregunta 4**. 

Escribe una función llamada `clasificar_porcentaje_cosechado` que reciba una fila de una tabla con datos agrícolas y clasifique la eficiencia de la cosecha de acuerdo con el porcentaje cosechado respecto al área sembrada.

#### **La función debe:**
1. Leer los valores `'areaSembradaHa'` y `'areaCosechadaHa'` de la fila.
2. Si el área sembrada es cero, retornar la cadena `'Sin dato'`.
3. Calcular el porcentaje como:

   $$
   \text{porcentaje} = \frac{\text{Área Cosechada}}{\text{Área Sembrada}} \times 100
   $$

4. Clasificar la fila según las siguientes reglas:

   - Si el porcentaje es **menor que 60%**, retornar `'Bajo'`.
   - Si el porcentaje está entre **60% y 90%**, retornar `'Medio'`.
   - Si el porcentaje es **mayor que 90%**, retornar `'Alto'`.



Una vez definida la función, aplícala a una tabla llamada `cafe_rendimiento` utilizando `apply` y añade una nueva columna llamada `'ClasificacionCosecha'` con los resultados.

Ejemplo de uso
```python
cafe_clasificacion = cafe_rendimiento.with_column('clasificacionCosecha', cafe_rendimiento.apply(clasificar_porcentaje_cosechado))
```
La tabla resultante cafe_clasificacion tendrá una nueva columna que clasifica cada fila como 'Bajo', 'Medio', 'Alto' o 'Sin dato', dependiendo del porcentaje cosechado.

```python

def clasificar_porcentaje_cosechado(fila):
    sembrada = fila.item('areaSembradaHa')
    cosechada = fila.item('areaCosechadaHa')
    
    if sembrada == 0:
        return 'Sin dato'
    porcentaje = (... / ...) * 100
    
    if porcentaje < ...:
        return ...
    
    elif ... < porcentaje <= ..:
        return ...
    else:
        return ...
cafe_clasificacion = cafe_rendimiento.with_column('clasificacionCosecha', cafe.apply(clasificar_porcentaje_cosechado))
cafe_clasificacion
```

```python
grader.check("q4")
```

```python
def graficar_promedio_anual_por_departamento(tabla, departamento, col):
    agrupada = (
        tabla
        .select("departamento", "año", col)
        .where("departamento", are.equal_to(departamento))
        .group("año", np.mean)
    )
    
    col_media = f"{col} mean"
    agrupada.scatter("año", col_media)
```

Analiza las siguientes gráficas y saca tus propias conclusiones para cada una.

```python
col1="areaSembradaHa"
col2= "areaCosechadaHa"
col3="produccionT"

graficar_promedio_anual_por_departamento(cafe_clasificacion, "ANTIOQUIA", col1)
graficar_promedio_anual_por_departamento(cafe_clasificacion, "ANTIOQUIA", col2)
graficar_promedio_anual_por_departamento(cafe_clasificacion, "ANTIOQUIA", col3)
```

```python
Conclusion1=""
```

```python
Conclusion2=""
```

```python
Conclusion3=""
```

## Iteraciones


"expandir" un ciclo `for` significa escribir manualmente todo el código que el ciclo ejecutaría. El resultado es un código que hace lo mismo que el ciclo, pero sin usar su estructura. Por ejemplo, el siguiente ciclo:

```python
for num in np.arange(3):
    print("El número es", num)


Puede ser expandido así:

```python
print("El número es", 0)
print("El número es", 1)
print("El número es", 2)


```

Expandir un ciclo for es una excelente forma de entender qué hace el ciclo en cada paso. En este ejercicio, practicarás cómo desenrollar un ciclo for.
En la siguiente pregunta, escribe código que haga exactamente lo mismo que el código dado, pero con los ciclos for expandidos. Es buena idea ejecutar tanto tu respuesta como el código original para verificar que el comportamiento sea el mismo. (Por supuesto, si el código incluye algún comportamiento aleatorio, el resultado puede variar ligeramente).

```python
print("¿Qué hace una abeja en el gimnasio?")
for i in range(3): 
    print("...")
    print("¡Zum-ba!")
    print()
print("Fin del entrenamiento. 🐝💪")
```

# Practica de ciclos

**Paseo Aleatorio del Borracho**

El **paseo del borracho**, también conocido como **paseo aleatorio**, es un concepto clásico en teoría de probabilidades y matemáticas. Se refiere al camino formado por una serie de pasos aleatorios, típicamente en un espacio o red. Imagina a un borracho que comienza en un punto fijo y da pasos al azar, hacia adelante o hacia atrás, cada uno con una probabilidad definida. 

Este modelo es utilizado para ilustrar cómo las fluctuaciones aleatorias pueden influir en el movimiento o posición de un objeto a lo largo del tiempo. Se aplica en varios campos como la física, la biología, la economía y otras disciplinas que estudian procesos y fenómenos estocásticos. Es una herramienta fundamental para entender la difusión y procesos de varianza aleatoria en distintas áreas de estudio.

En los sigientes puntos, realizaremos su simulación.

Simula un paseo aleatorio de un borracho que realiza 100 pasos. Cada paso puede ser hacia adelante (+1) o hacia atrás (-1) con igual probabilidad. Ejecuta esta simulación 10,000 veces para analizar el desplazamiento final del borracho, para ello realiza lo siguiente :



1. **Define una función llamada `simulacion_paseo_borracho`**  
   Esta función debe generar 100 pasos aleatorios, cada uno con valor +1 o −1 (con igual probabilidad), y devolver la suma total de esos pasos, es decir, la posición final alcanzada.

2. **Crea un arreglo vacío llamado `desplazamientos_finales`**  
   Este arreglo servirá para almacenar la posición final del borracho en cada una de las simulaciones.

3. **Simula el experimento 10,000 veces**  
   Usa un ciclo `for` con `np.arange` para repetir la simulación 10,000 veces.  
   En cada iteración:
   - Llama a la función `simulacion_paseo_borracho()`.
   - Agrega el valor retornado al arreglo `desplazamientos_finales`.

4. **Construye una tabla llamada `sim`**  
   Esta tabla debe tener dos columnas:
   - `"desplazamiento"`: contiene las posiciones finales de cada recorrido.
   - `"simulacion"`: contiene los índices de simulación (del 1 al 10,000).

5. **Muestra las primeras 10 filas de la tabla** usando `.show(10)` para verificar que todo funcione correctamente.

#### 💡 Nota
Puedes usar `np.random.choice([-1, 1], size=1000)` para generar los pasos aleatorios, y `np.sum()` para obtener la posición final.

---


```python
def simulacion_paseo_borracho():
    ...
    return ...

num_repeticiones = 10000
desplazamientos_finales = make_array()

for _ in np.arange(num_repeticiones):
    nuevo_desplazamiento = simulacion_paseo_borracho()
    desplazamientos_finales = np.append(...., nuevo_desplazamiento)

simulacion = Table().with_columns(
    "...", desplazamientos_finales,
    "....", np.arange(1, num_repeticiones + 1)
)

```

```python
grader.check("q5")
```

```python
sim.hist("desplazamiento", bins=range(-100, 100,2))
```

¿Qué puede concluir del histograma?, interpreta muy bien cada una de las variables

```python
conclusion="Escribe aqui tu respuesta"
```

¡Terminó la tarea 5, recuerda que NO tuvimos tarea 4!

**Información sobre la entrega:** Asegúrese de correr todas las pruebas y verificar que se ejectuen correctamente. Aquí, seleccione **Guardar Notebook** En el menú de **Archivo** y **corra la  última celda**. Descargue el archivo .zip y el archivo con extensión .ipynb y subálo al enlace compartido en la página del curso para su entrega. El nombre de esta tarea es "HW 03 ". **Es su responsabilidad asegurarse que su trabajo se guarde antes de correr la última celda**.



`Referencia`: Este material esta basado en [Data 8: The Foundations of Data Science](https://www.data8.org/) - University of California, Berkeley.

Estimad@ estudiante de Fundamentos de Ciencia de Datos

¡Felicitaciones por terminar la tarea 5!

<img src="./Casafus.jpg" width="50%" alt="Close up picture of dog smiling."/>

Perro de la semana



## Entrega

Asegúrese que ha ejecutado todas las celdas en su notebook en orden antes de ejecutar la celda de abajo. De esta forma todas las imágenes/gráficas apareceran en la salida. La celda de abajo va a generar un archivo con extensión .zip para que usted lo entregue. **¡Por favor guarde antes de exportar!**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> assert cafe.labels == ('departamento', 'municipio', 'año', 'periodo', 'areaSembradaHa', 'areaCosechadaHa', 'produccionT', 'cicloDeCultivo')
```


Tests for q2 (points: 1)
```
>>> assert np.isclose(cafe_rendimiento.column('rendimiento').mean(), 0.9198174805793192, 0.01)
```


Tests for q3 (points: 1)
```
>>> assert np.isclose(np.mean(produccion_departamento(cafe_rendimiento, 'ANTIOQUIA').column('produccionTotal')), 116863.33333, 0.1)
```


Tests for q4 (points: 1)
```
>>> b = np.count_nonzero(cafe_clasificacion.column('clasificacionCosecha') == 'Bajo')
>>> m = np.count_nonzero(cafe_clasificacion.column('clasificacionCosecha') == 'Medio')
>>> a = np.count_nonzero(cafe_clasificacion.column('clasificacionCosecha') == 'Alto')
>>> assert np.isclose((462, 4449, 2350) == (b, m, a), True)
```


Tests for q5 (points: 1)
```
>>> assert np.isclose(round(np.mean(sim.column('desplazamiento'))), 0.0, 0.1)
```
